---
title: 'Get Acquainted with the DIG dataset'
description: 'In this chapter you''ll use the tidyverse instruments to explore and get familiar with the dataset'
---

## An exercise title written in sentence case

```yaml
type: NormalExercise
key: 7e870be402
lang: r
xp: 100
skills: 1
```

In 1997 this paper came out on the prestigious NEJM. The aim was to compare the use of cardiac glycosides in the treatment of patients with heart failure and normal sinus rhythm, since it is controversial.

`@instructions`
- Load the Tidyverse package
- Read the file dig.csv using the appropriate function
- use glimpse() to look at the structure

`@hint`
- Here is the hint for this setup problem. 
- It should get students 50% of the way to the correct answer.
- So don't provide the answer, but don't just reiterate the instructions.
- Typically one hint per instruction is a sensible amount.

`@pre_exercise_code`
```{r}
# Load datasets and packages here.
```

`@sample_code`
```{r}
#load the Tidyverse package
library(___)

#read the dig.csv file with the appropriate function from the package
dig <- ___()

#explore the content of the dataset
___

```

`@solution`
```{r}
# Answer goes here
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```

`@sct`
```{r}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```
